---
title: Template Chapter 1
description: This is a template chapter.

--- type:NormalExercise lang:r xp:10 skills:1 key:a57baa03e8
##Beispiel: Vektoren


bla

*** =instructions

bla

*** =hint 


xxx

*** =pre_exercise_code

```{r}
library(dplyr)
```


*** =sample_code

```{r}
# alle benötigten Objekte sind bereits eingelesen
daten_auswahl1 <- ___

```

*** =solution

```{r}
daten_auswahl1 <- 1
```

*** =sct

```{r}
test_object("daten_auswahl1") 
test_error()
success_msg("Sehr gut!")
```
